// Write a program to print a series of number starting from 0 in which a number corresponds to sum of previous two nunmbers.
// Use function.
// For hint type "hint"

int specialseries(int i) {

   // code

int  main() {

   int i;
	
   for (i = 0; i < 10; i++) {
      printf("%d\t\n", specialseries(i));
   }
	
   return 0;
}

// I AM NOT DONE!