// Jessica handed over the charge of poster making to her assistant Nerrisa.
// Now she wants to get the names that are there in "cigcomponent"
// Help her read the existing file and get the names.
// For hints type "hint"

#include<stdio.h>
#include<stdlib.h>
#include<string.h>

int main()
{
    FILE*fptr;
    char str[100];
    int i=0;
    
    // open file

    while(fscanf(fptr,"%s",str)!=EOF)
    {
        // copy names to an array
        
    }
    
    
        printf("\nThe array is ; [ ");
        for(int j=0;j<i;j++)
        {
            printf("'%s",arr[j]);
        }
        
        
        // close file


        return 0;
}

// I AM NOT DONE!