// Students have got their numbers from the given values in the previous program. Now they want to calculate their aggregate.
// Help them to do it using for loop.
// For hint type "hint"
#include <stdio.h>
int main()
{
    int NumSub, marks, aggregate = 0;
    printf("enter total number of subjects\n");
    scanf("%d", &NumSub);
    printf("enter marks in subjects:\n");
    for (...........) //enter conditions

    {
        //scan individual marks and aggregate
    }
    printf("aggregate of marks is %d\n", aggregate);

    //I AM NOT DONE!