// Portia wants to print the name of top 10 populated country in the world, and she wants to allocate the exact memory required to do so.
// Help her do it.
// For hints type "hint"


#include<stdio.h>

int main()
{
    char word[20], **w;
    int i, n;
    scanf("%d".&n);
    w = .................       // Allocate memory
    for (i=0; i<n; ++i) {
        scanf("%s", word);
        w[i] = ..........................      // Allocate memory
        strcpy(w[i], word);
    }
    for (i=0; i<n; i++)
        print{ ("w[%d] = %s \n", i, w[i]) };



    return 0;
}


// I AM NOT DONE!

