// Mr.Ariel, the class teacher of class X has made the marksheet of students in a strange manner. He has put certain alphabets instead of integers in place of marks.
// He has asked students to crack the value of the alphabets to get their marks.
// Can you do it?
// For hints type "hint"

#include <stdio.h>

int main()
{
    datatype englishmarks = 'd';
    datatype mathsmarks = 'X';
    datatype sciencemarks = 'V';
    datatype HCGmarks = 'a';
    datatype artmarks = 'Z'
    printf("the marks are % ,% ,% ,% ,% \n", englishmarks, mathsmarks, sciencemarks, HCGmarks,artmarks);
    return 0;
}
// I AM NOT DONE!