// Miranda has drawn a closed structures consisting of three sides. Now she wants to calculate its area using macro.
// Help her to do it.
// For hint type "hint"

#include<stdio.h>

// define macro

int main()
{
	float base, height, a;
	printf("enter base of triangle: ");
	scanf("%f",&base);
	
	printf("enter height of the triangle: ");
	scanf("%f",&height);
        a=area(b,h); 
	printf("Area of the triangle : %f\n",a);
	return 0;
}

// I AM NOT DONE!