// Form an enum of transgenic plants namely Bt cotton, wheat, brassica napus, golden rice, Flavr Savr Tomato and print their values.
// Initialise the value to 1
// For hints type "hint"

#include<stdio.h>

// declare enum

int main (){

    enum trans_plants plant;


    // print values

    return 0;
}

// I AM NOT DONE!