// Ferdinand is a teacher at a kindergarten school.
// He wants to teach his students the name of the month and the number at which they occur, help him do that using switch and enumerator.
// For hints type "hint"


#include<stdio.h>
// declare the enum months
int main()
{
    enum months number;
    int n;
    printf("enter no. 1-12\n");
    scanf("%d",&n);
    switch(n)
    {
        
    case 1:
        printf("January");
        break;
    .
    .
    .
    .
    .
    .
    .   // write for each month
    .
    .
    .
    .
    .
    .
    .
    .

    default :
        printf("invalid number!!!");
    }
}

// I AM NOT DONE!