// Solve the question structure1 using pointers.
// For hints type "hint"

#include<stdio.h>
struct captain{

    //declare the structure

}
int main(){
    struct student red, yellow, green, blue;

    // declare pointer to the structure

    // access the elements and print the value

    return 0;

}


// I AM NOT DONE!