
// Mr.A needs the list of population percentage having RhD present in their blood group i.e having positive blood groups.
// He has got the data in which the sign before the number denotes whether the group is negative or positive.
// A=+33%, A=-7%, B=+9%, B=-2%, O=+37%, O=-8%, AB=+3%, AB=-1%.
// Help him print the values of RhD positive blood population using the datatype that prints only positive values.
// For hints type "hint"
// All the best!

#include <stdio.h>
int main()
{
    datatype A = +33;
    datatype B = +9;
    datatype O = +37;
    datatype AB = +3;

    printf(" The percentage of population having A+ are %, B+ are %, O+ are %, AB+ are %", A, B, O, AB); //put correct format specifier
}

// I AM NOT DONE!