// Learnt from his previous experience Bassanio tried to implement the idea to get the accurate contribution of his individual friends in overall profit of their shared business.
// But unfortunately he faced errors again. Help him to solve the issue.
// For hint type "hint" 



#include<stdio.h>
int main (){
    float Salanio = 5.1235455486;
    float Salerio = 5.12354545887787;
    float Gratiano = 2.321478569458213;
    printf("%\n",Salanio);
    printf("%\n",Salerio);
    printf("%\n", Gratiano);
    return 0;
}

//I AM NOT DONE!