// Students need only the aggregate of marks, so declaraing each subject marks individually and then adding them is wastage of time.
// Find a more effective way which takes only the marks as input of an specific data structure and serves the purpose.
// For hint type "hint"

#include <stdio.h>

int main()
{
    XXXXX; //declare the data structure required

    int NumSub;
    int i;
    int aggregate = 0;
    printf("enter the number of subjects: \n");
    scanf("%d\n", &NumSub);

    for (........................)
    {

        //scan the required data structure that holds the marks obtained
    }

    printf("\nPrinting marks obtained: \n\n");

    for (...........................)
    {
        // print marks obtained
    }

    for (........................)
    {

        //calculate the aggregate
    }
    printf("\nTotal marks obatained is : %d\n", aggregate);

    return 0;
}
//I AM NOT DONE!